:orphan:

Problem Set 2 (due midnight Jan 22)
===================================

..  Copyright (C) Paul Resnick.  Permission is granted to copy, distribute
    and/or modify this document under the terms of the GNU Free Documentation
    License, Version 1.3 or any later version published by the Free Software
    Foundation; with Invariant Sections being Forward, Prefaces, and
    Contributor List, no Front-Cover Texts, and no Back-Cover Texts.  A copy of
    the license is included in the section entitled "GNU Free Documentation
    License".


.. assignment::
  :name: PS2
  :assignment_type: problem_set
  :questions: p2_2_1 6, p2_2_2 10, ps_2_3 10, ps_2_4 5, ps_2_5 5, ps_2_6 7, ps_2_7 7
  :deadline: 2017-01-23 04:00
  :points: 50
  :autograde: unittest

.. activecode:: p2_2_1
    :language: python
    :autograde: unittest
    :hidecode:

    **1.** Write code to assign the FIRST CHARACTER OF EVERY WORD in the english sentence ``sentence`` to the string ``first_characters``.
    ~~~~
    sentence = "The quick brown fox jumps over the lazy dog."

    # assign the string first_characters to the first character of every word in sentence
    # DO NOT HARD CODE your answer


    =====

    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

        def testOne(self):
           self.assertEqual(first_characters, 'Tqbfjotld', "Testing that first_characters has been set to first character of every word in sentence")
           self.assertNotIn('Tqbfjotld', self.getEditorText(), "Testing that you did not hard code your answer. (Don't worry about Actual and Expected Values.)")

    myTests().main()

.. activecode:: p2_2_2
    :language: python
    :autograde: unittest
    :hidecode:

    **2.** Assign the variable ``fl`` the value of the first element of the string value in ``original_str``. Use string indexing to assign the variable ``last_l`` the value of the last element of the string value in ``original_str``. Write code so that will work no matter how long ``original_str``'s value is.
    ~~~~
    original_str = "The quick brown rhino jumped over the extremely lazy fox."

    # assign variables as specified below this line!

    =====

    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

        def testOne(self):
           self.assertEqual(fl, original_str[0], "Testing that fl has been set to first char in original_str")
           self.assertEqual(last_l, original_str[-1], "Testing that last_l has been set to last char in original_str")
           self.assertIn('-1', self.getEditorText(), "Testing that you indexed correctly. (Don't worry about Actual and Expected Values.)")

    myTests().main()

.. activecode:: ps_2_3
    :include: addl_functions_2
    :language: python
    :hidecode:

    **3.** There is a function we are giving you for this problem set that takes two strings as inputs, and returns the length of both of those strings added together, called ``add_lengths``. We are also including the functions from Problem Set 1 called ``random_digit`` and ``square`` in this problem set.

    Now, take a look at the following code and related questions, in this code window.
    ~~~~
    new_str = "'Twas brillig"

    y = add_lengths("receipt","receive")

    x = random_digit()

    z = new_str.find('b')

    l = new_str.find("'")

    # notice that this line of code is made up of a lot of different expressions
    fin_value = square(len(new_str)) + (z - l) + (x * random_digit())

    # DO NOT CHANGE ANY CODE ABOVE THIS LINE
    # But below here, putting print statements and running the code may help you!

    # The following questions are based on that code. All refer to the types of the
    #variables and/or expressions after the above code is run.

    #####################

    # Write a comment explaining each of the following, after each question.
    # Don't forget to press **run** to save!

    # What is square?

    # What type of object does the expression square(len(new_str)) evaluate to?

    # What type is z?

    # What type is l?

    # What type is the expression z-l?

    # What type is x?

    # What type is y?

    # What is random_digit? How many inputs does it take?

    # What type does the expression x * random_digit() evaluate to?

    # Given all this information, what type will fin_value hold once all this code is run?

    ====

    print "\n\nThere are no tests for this problem"


.. activecode:: addl_functions_2
    :nopre:
    :hidecode:

    def square(num):
        return num**2

    def greeting(st):
        #st = str(st) # just in case
        return "Hello, " + st

    def random_digit():
        import random
        return random.choice([0,1,2,3,4,5,6,7,8,9])

    def add_lengths(str1, str2):
        return len(str1) + len(str2)

.. activecode:: ps_2_4
   :language: python
   :autograde: unittest
   :hidecode:

   **4.** Write code that uses iteration to print out each element of the list ``several_things``. Then, write code to print out the TYPE of each element of the list called ``several_things``. Note that these should be in separate loops.
   ~~~~
   several_things = ["hello", 2, 4, 6.0, 7.5, 234352354, "the end", "", 99]

   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

     def test_output(self):
         self.assertIn('for', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
         self.assertIn("<type 'str'>\n<type 'int'>\n<type 'int'>\n<type 'float'>\n<type 'float'>\n<type 'int'>\n<type 'str'>\n<type 'str'>\n<type 'int'>", self.getOutput(), "Testing output (Don't worry about actual and expected values).")

   myTests().main()

.. activecode:: ps_2_5
   :language: python
   :autograde: unittest
   :hidecode:

   **5.** The code provided does not iterate over the words in the English sentence that's stored in the variable ``sent``. Why not? Write a comment in the box below explaining why not. (Hint: Knowing what you know about how computers and programming languages deal with sequences, what do you need to do to make sure you can iterate over the words in the sentence?)

   Then, write code that assigns a variable word_list to hold a LIST of all the WORDS in the string sent. (It's fine if words include punctuation.)
   ~~~~
   sent = "The magical mystery tour is waiting to take you away."

   for x in sent:
      print x

   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

      def testOne(self):
         print "No tests for the comment -- we have to read those!\n"
         self.assertEqual(word_list, sent.split(), "Testing that word_list has been set to a list of all the words in sent")

   myTests().main()

.. activecode:: ps_2_6
   :language: python
   :autograde: unittest
   :hidecode:

   **6.** Write code to store the LENGTH of every word in the English sentence that's stored in the variable ``sent`` into the variable ``word_lengths``. ``word_lengths`` should be a list where the element at any index i is the length of word i.

   ~~~~
   sent = "The magical mystery tour is waiting to take you away"


   =====

   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

      def testOne(self):
         print "No tests for the moment -- we have to read those!\n"
         self.assertEqual(word_lengths, map(len, sent.split()), "Testing that word_lengths is equal to the lenth of every word in sent")

   myTests().main()

.. activecode:: ps_2_7
    :language: python
    :hidecode:


    **7.**
    * Read :ref:`Object Instances and Turtle graphics<turtles_chap>`

    Write a program that uses the turtle module to draw something interesting using a ``for`` loop. It doesn't have to be complicated, but draw something different than we did in the textbook or in class. (Optional but encouraged: post a screenshot of the artistic outcome to the Facebook group, or a short video of the drawing as it is created.) (Hint: if you are drawing something complicated, it could get tedious to watch it draw over and over. Try setting ``.speed(10)`` for the turtle to draw fast, or ``.speed(0)`` for it to draw super fast with no animation.)
    ~~~~
    import turtle